# Homework 2: Scripting and Jobs (Bash, cnt'd)
Curated by Sabah Ul-Hasan, Huitian Dao, Karthik Gangavarapu, and Andrew Su

Homeworks are due every Tuesday at 8 AM PST.<br>
Download your homework as a notebook for TAs to provide feedback: HW#_LastName.ipynb<br>
Make sure to save your work constantly so the efforts don't get lost!

In this assignment, you will <br>
A2.aa Refresh your understanding of bash<br>
A2.ab Explore pipes in-depth (+awk) <br>
A2.ba Familiarize yourself with loops and Git <br>
A2.bb Compile what you've learned as job to run on the cluster (HPC)<br>

Note that there is a section at the very end (A2.bb) to point out any enhancements you personally would make to this assignment if you were to lead the course. Feel free to give it a try!

### A2.aa: Refresh understanding of bash [20-30 min] Group of 4 
1. [5 min] In your group, identify four items you each considered to be the most difficult part of Homework 1. Then, break into pairs based on who in the group found it easier to navigate vs harder to navigate. In other words, pair yourself with someone who better understood what you found difficult and vice versa. Denote what these sections are below as a,b,c,or d.
2. [10 min] Spend time discussing your troubleshooting process of working through what the other person found as difficult with one another. Write down what you learned. What was the before vs after result?
3. [15 min] Now, share as a whole group. Write down what you gained from the discussion in how you will be approaching bash scripting differently than you would do before.

```python
# 1a
# 1b
# 1c
# 1d
```

```python
# 2
```

```python
# 3
```

### A2.aa: Explore pipes in-depth (+awk) [20-30 min] TPS 
[10-15 min]: On your own time, try the following exercises

### Question 1 (1 point)
Describe your current understanding of __pipes__. What are the various approaches, where does `awk` come into play?

<!-- grade_id: pipe_answer, points: 1, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===  
__Pipe__ is used to connect STDOUT to STDIN.  
`awk` can be used to process text files, especially delimited files based on seperated fields.  
=== END MARK SCHEME ===  

### Question 2 (1 point)
Find the error and fix to run the script below. Describe the error in a sentence.  

```
%%bash
awk '{print $3}' data/new_folder/artists_new.txt
```

What is `{print $3}` doing? Describe what `awk` is doing this, and its limitations in recognizing delimiters (Hint: tab). 

<!-- grade_id: fix_awk_answer, points: 1, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===  
Need to set the field seperator as tab  
=== END MARK SCHEME ===  

<!-- grade_id: fix_awk_code, points: 0, grade: false, solution: true -->
```python
%%bash --out var2

### BEGIN MY SOLUTION
awk -F "\t" '{print $3}' data/new_folder/artists_new.txt
### END MY SOLUTION
```

<!-- grade_id: fix_awk_test, points: 1, grade: true, solution: false -->
```python
# check the first line of the output
assert var2.split('\n')[0] == 'Baroque'
# check the last line of the output
assert var2.split('\n')[-2] == 'Early Renaissance'
### BEGIN HIDDEN TESTS
assert len(var2.split('\n')) == 151
value_counts = dict((x, var2.split('\n').count(x)) for x in set(var2.split('\n')))
assert value_counts == {'': 1,
 'Romanticism': 18,
 'Neoclassicism': 10,
 'High Renaissance': 5,
 'Early Renaissance': 12,
 'Mannerism': 10,
 'Medieval': 8,
 'Northern Renaissance': 6,
 'Baroque': 61,
 'Rococo': 10,
 'Realism': 9,
 'Impressionism': 1}
### END HIDDEN TESTS
```

### Question 3 (1 point)
Use `awk`, `sort`, and/or `uniq` to find the most common last names among all artists in 'artists.txt' [10-15 min]  
Now, discuss your findings with your partner. What was similar, and what was different?  

<!-- grade_id: awk_sort_uniq_answer, points: 1, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===  
      5 KOBELL
      5 KNIP
=== END MARK SCHEME ===  

<!-- grade_id: awk_sort_uniq_code, points: 0, grade: false, solution: true -->
```python
%%bash

### BEGIN MY SOLUTION
awk -F ",| " '{print $1}' data/new_folder/artists_new.txt | uniq -c | sort -r
### END MY SOLUTION
```

On your own: 

### Question 4 (1 point)
Now that we've worked through these in pairs, create a bash script that integrates two items we haven't yet tried in the exercises, such as NF, to answer a question about the artists.txt file. Describe briefly what this does. Have this complete pior to next class.

<!-- grade_id: integrate_answer, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===  
Example: Using different field seperators, find out how many rows have certain number of columns  
=== END MARK SCHEME ===  

<!-- grade_id: integrate_code, points: 0, grade: false, solution: true -->
```python
%%bash

### BEGIN MY SOLUTION
echo "by tab"
awk -F '\t' '{print NF}' data/new_folder/artists_new.txt | sort | uniq -c
echo "by comma"
awk -F ',' '{print NF}' data/new_folder/artists_new.txt | sort | uniq -c
### END MY SOLUTION
```

Output:
```
by tab
    150 4
by comma
     10 1
    128 2
     12 3
```

### Question 5 (1 point)
Reflections: On an integer scale of `1-5`, `1` being very easy and `5` being the most difficult, where would you rank today's exercises?  
Provide brief notes on what was easy vs difficult, and why.

Install Git and review prior to next class: https://git-scm.com/book/en/v2/Getting-Started-Installing-Git

<!-- grade_id: reflections_answer, points: 1, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===  

=== END MARK SCHEME ===  

### A2.ba: Loops and Git [20-30 min] Group of 3
#### ba.i Loops <br>
[10 min] On your own
1. Describe how a loop operates. Why and when do you want to use a loop, vs a single command string?
2. Write a bash for loop that outpouts letters a-e
3. What are the different types of loops, and when might you want to try one over another? Script an example.

#### ba.ii Git <br>
[15-20 min] Discuss and execute the following with your group -- be careful of permanently changing something!
4. Save one of the scripts you've created and push it to a new Github repository connected to your own account labeled as "Applied-Bioinformatics_Course". Do this as a group through one account, and then affirm your understanding by ensuring everyone has their own script uploaded in their repository. Use the terminal to do this, not the website (though you can confirm it pops up on the website)! Share the final URL below.

```python
# Loop:
# A sequence of instructions to be repeated until a certain condition is reached
# Why: to conveniently repeat instructions
# When: when an instruction need to be executed repeatedly and the condition for ending executioncan be described
```

```python
%%bash
# 2
for i in {a..e}
do
  echo $i
done
```

Output:
```
a
b
c
d
e
```

```python
%%bash
# 3
# For / While / Util

for i in {1..3}
do
  echo $i
done
```

Output:
```
1
2
3
```

```python
%%bash
i=0

while [ $i -lt 3 ]
do
  echo $i
  i=$((i+1))
done
```

Output:
```
0
1
2
```

```python
%%bash
i=0

until [ $i -gt 3 ]
do
  echo $i
  i=$((i+1))
done
```

Output:
```
0
1
2
3
```

```python
# 4
# git add -A
# git commit "NA"
# git push
```

### A2.bb: Create a script to run a job on the cluster (HPC) [20-30 min] Individual
Get excited, we're about to apply what we've learned from a technical perspective to the science-y stuff! Familiar with 'fasta files' from NCBI? Here we go!
https://blast.ncbi.nlm.nih.gov/Blast.cgi?CMD=Web&PAGE_TYPE=BlastDocs&DOC_TYPE=BlastHelp

#### bb.i Applying what we've learned on bash to biology <br>
1. In your "Applied-Bioinformatics_Course" Github repository, create a folder titled "HW2_[Lastname]"<br>
*Always check that you're in the correct working directory
2. Create a bash loop that renames all the files in the folder data/fastq folder from .fq to .fastq<br>
*Have all these files copied and moved into the 'new_folder' directory so as not to overwrite the original .fq
3. Print all nucleotide sequences from ./data/fastq/simulated1.R1.fastq by replacing all 'T's in the nucleotide sequences with 'U's

#### bb.ii Checking that we can access the HPC *make sure you're set up for VPN <br>
4. On your terminal, are you able to login to the TSRI cluster? Y/N <br>
```ssh [yourlogin]@login01.scripps.edu``` <br>
[Type out any obstacles you encounter here] <br>
5. What is meant by "HPC", and why do we need one vs our own personal computer or Github? 

HPC User Manual: http://docs.adaptivecomputing.com/torque/4-0-2/help.htm#topics/2-jobs/submittingManagingJobs.htm#topics/2-jobs/submittingManagingJobs.htm?TocPath=2.0%20Submitting%20and%20managing%20jobs%7C_____0

#### bb.iii Running a job on the HPC to explore a biological question On Checking that we can access the HPC <br>
6. Use scp to download the file 'models.txt' from <br> 
/gpfs/home/hdiao/Applied-Bioinformatics-Course/week1-2 <br>
And append "Drosophila melanogaster" into the file
7. Log in to the HPC, copy the folder<br>
/gpfs/home/hdiao/Applied-Bioinformatics-Course/week1-2<br>
to your HPC home directory
8. Run fastqc for the file pairend_fastq/SRR396637.sra_1.fastq. Submit a job, check the job process, archive and zip the output files.
9. Change permission for the file hs_TBX21.fasta so that only owner can read and write it.
10. Reflections: On an integer scale of 1-5, 1 being very easy and 5 being the most difficult, where would you rank the A2.bb assignment and why? Include some notes about what you gained from this past week that you didn't know before as well as notes of things you already knew. 
11. Improvements: Notice a mistake, or an area of this homework that could have been improved? Include your thoughts here.

```python
%%bash
# 2

cp -R data/fastq data/new_folder
pwd
cd data/new_folder/fastq

for i in *.fq
do
  mv $i ${i/.fq/.fastq} 
done
```

Output:
```
/Users/sulhasan/Downloads/HW-Keys
```

```python
%%bash
# 3

cat data/new_folder/fastq/simulated1.R1.fastq | awk "NR%4==2" | tr T U
```

Output:
```
UGGGAAAACCAGGAGAGUUCUUCCUGAAAUAGUCCGUGAAGCCAUAAAAAGAAGACUCCGUACUGUGAUC
ACUUACCAUUUGUCAUGGCCCUGGGACUAACCGCUGUGAGGCUGGUCGACCCGAUCAACGUGGUGGGACU
GCCAGACAACCCCGUUUAUUAGAGAGGACGCUGACCCUUGUGUGGGGGCCUCAUAGCUUCCAUGGUAAGC
CUCUAAGCUGAUGCUGGAACUUGAUCCACCAUUUGGGGACUCUUACAUUGUCAUAGGAGUCGGGGAGAAG
GGCAUCACCGAGGUGUGCAGAGAAGAGGCCCGCCGCGCCCUCAAGGACGGUGUGGCAACGGGAGGCCAUG
GAGAUUGCCCGUGCCUGUGAACGAGCUGCCCCACGGGUGGAAGGCUUUGGGGAAAUCGUAUUUCGUCAGA
AUACCAUCAUCCCAGGCUUGCUUGACUGCUGCUGGCAAUCUACGGGGGGAGUCAGGAUGGUACUUGUACC
UUCUUCUUCACACUGCCUUUUCCCUUCAGAGAGAGGAGCAAAAACCCCCCGCAAGUAGCAAGGCCUGCUC
GAAAGAUCUGGCUGCCAAGCUGAGAAUAAUCAAUGCUAGGAAGGAGAAGAAGAGACGAGGCGCAGAUACU
GUUUUCAAGGAAAAAGUGGACACUAGGGUGCCAGACCCCCAAGAAGGCACUCGUCAGGUUAUGAGCAUGG
ACAUGCAAACUCUGACCCCAGUUGGGAGGUUGAUAACCGCUAACCCCGUACUCACUGAAAGCACUGAGAA
GGCUUGAAAAAAGACCAGGAGCCUUUUGUAUAAAAGUGUUGUGCCCAUACACCAGCACUAUGCUGGAAAC
UCUUUUCUCUUUUUCCCAUCAUGUUGUACACACAGCUCUGGCACUCUCCUCUCAGGUGGUGCACUCUUUG
CACUGAUUUGAAAGGUGCUCCAAAAAUUUGAUGGAUGCGCUUGCCCAAUGAGUUGAGAGCGCCUCCAACU
CUUUUC
```

```python
# HPC: high performance computing (higher computational capacity, for solving large scale problems)
```

```python
%%bash
# 4

echo "Drosophila melanogaster" >> models.txt
```

```
#!usr/bin/env bash

module load fastqc

fastqc pairend_fastq/SRR396637.sra_1.fastq

```

```python
%%bash

chmod 700 hs_TBX21.fasta
```

Output:
```
chmod: hs_TBX21.fasta: No such file or directory
~/opt/anaconda3/lib/python3.8/site-packages/IPython/core/magic.py in <lambda>(f, *a, **k)
    185     # but it's overkill for just that one bit of state.
    186     def magic_deco(arg):
--> 187         call = lambda f, *a, **k: f(*a, **k)
    188 
    189         if callable(arg):

~/opt/anaconda3/lib/python3.8/site-packages/IPython/core/magics/script.py in shebang(self, line, cell)
    243             sys.stderr.flush()
    244         if args.raise_error and p.returncode!=0:
--> 245             raise CalledProcessError(p.returncode, cell, output=out, stderr=err)
    246 
    247     def _run_script(self, p, cell, to_close):

CalledProcessError: Command 'b'\nchmod 700 hs_TBX21.fasta\n'' returned non-zero exit status 1.
```

```python

```